---
title: Test
description: Test
---

## Conditional Formatting a Date Range

```yaml
type: NormalExercise
key: c909bf2ad1
lang: python
xp: 100
skills: 2
```

Conditional formatting can bring attention to data that occurs within a certain _time period_. Working with dates in spreadsheets is more difficult than working with plain numbers, but Sheets has options available in the conditional formatting menu that will make this process easier.

`@instructions`
- Check the dates to ensure they are valid. Fix dates as needed.
- Select and conditional format dates in the last week.
- Tell us what item has been selling the best in the last week.
- Select and conditional format dates from October 10, 2018 through October 13, 2018.

`@hint`
- Use the DATEVALUE function to check for valid dates.

`@pre_exercise_code`
```{python}
# Load datasets and packages here.
```

`@sample_code`
```{python}
# Your
# sample
# code
# should
# be
# ideally
# 10 lines or less,
# with a max
# of 16 lines.
```

`@solution`
```{python}
# Answer goes here
# Make sure to match the comments with your sample code
# to help students see the differences from solution
# to given.
```

`@sct`
```{python}
# Update this to something more informative.
success_msg("Some praise! Then reinforce a learning objective from the exercise.")
```

---

## Conditional Formatting with Dates

```yaml
type: VideoExercise
key: 165b42fd50
xp: 50
```

`@projector_key`
7cd445d00011c3bbcbfdf671e32f0caa
