---
title: Test
description: Test
---

## Conditional Formatting a Date Range

```yaml
type: NormalExercise
key: c909bf2ad1
lang: python
xp: 100
skills: 2
```

Conditional formatting can be used to bring attention to data that occurs within a certain _time period_. Working with dates in spreadsheets is more difficult than plain numbers, but Sheets has options available in the conditional formatting menu that will make this process easier.

`@instructions`
- Check the dates to ensure they are valid. Fix dates as needed.
- Move the date column next to the item column.
- Select and conditional format dates in the month of October.
- Tell us what item sells best in October.

`@hint`
- Here is the hint for this setup problem. 
- It should get students 50% of the way to the correct answer.
- So don't provide the answer, but don't just reiterate the instructions.
- Typically one hint per instruction is a sensible amount.

`@pre_exercise_code`
```{python}
# Load datasets and packages here.
```

`@sample_code`
```{python}
# Your
# sample
# code
# should
# be
# ideally
# 10 lines or less,
# with a max
# of 16 lines.
```

`@solution`
```{python}
# Answer goes here
# Make sure to match the comments with your sample code
# to help students see the differences from solution
# to given.
```

`@sct`
```{python}
# Update this to something more informative.
success_msg("Some praise! Then reinforce a learning objective from the exercise.")
```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 165b42fd50
xp: 50
```

`@projector_key`
7cd445d00011c3bbcbfdf671e32f0caa
